---
title: 'Evaluating your model'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

How do you evaluate your model?

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Benchmarking your model

```yaml
type: VideoExercise
key: 266003694c
xp: 50
```

`@projector_key`
74af834eba04b5f46dc2ffba5d7d6ec2

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 8e3d3a581f
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`
